1. **Restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart)
2. **Run all cells** (in the menubar, select Cell$\rightarrow$Run All).
3. __Use the__ `Validate` __button in the Assignments tab before submitting__.

__Include comments, derivations, explanations, graphs, etc.__ 

You __work in groups__ (= 3 people). __Write the full name and S/U-number of all team members!__

```python
* Student 1 Stian Grønlund, s1122151:
* Student 2 name, S/U-number:
* Student 3 name, S/U-number:
```

---

# Assignment 1 (Statistical Machine Learning 2024)
# **Deadline: 27 September 2024**

## Instructions
* Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE` __including comments, derivations, explanations, graphs, etc.__ 
Elements and/or intermediate steps required to derive the answer have to be in the report. If an exercise requires coding, explain briefly what the code does (in comments). All figures should have titles (descriptions), axis labels, and legends.
* Please do __not add new cells__ to the notebook, try to write the answers only in the provided cells. Before you turn the assignment in, make sure everything runs as expected.
* __Use the variable names given in the exercises__, do not assign your own variable names. 
* __Only one team member needs to upload the solutions__. This can be done under the Assignments tab, where you fetched the assignments, and where you can also validate your submissions. Please do not change the filenames of the individual Jupyter notebooks.

For any problems or questions regarding the assignments, ask during the tutorial or send an email to charlotte.cambiervannooten@ru.nl and janneke.verbeek@ru.nl .

## Introduction
Assignment 1 consists of:
1. Polynomial curve fitting (50 points);
2. __Gradient descent (25 points);__
3. Fruit boxes (25 points);
4. Probability factorization (BONUS 10 points);

## Libraries

Please __avoid installing new packages__, unless really necessary.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it to at least version 3."

# Necessary imports (for solutions)
import math
import numpy as np
import matplotlib.pyplot as plt
from collections import namedtuple

# Set fixed random seed for reproducibility
np.random.seed(2022)
```

## Gradient descent (weight 25)
In this exercise, we consider the gradient descent algorithm for function minimization. When the function to be minimized is $E(\mathbf{x})$, the gradient descent iteration is  
\begin{equation*}
\mathbf{x}_{n+1} = \mathbf{x}_n - \eta \nabla E(\mathbf{x}_n) \tag{6}
\end{equation*}
where $\eta>0$ is the so-called learning-rate. In the following, we will apply gradient descent to the function
\begin{equation*}
h(x,y) = 100(y - x^2)^2 +(1 - x)^2 \label{banana} \tag{7}
\end{equation*}
### Exercise 2.1
Make a plot of the function $h$ over the interval $[-2 \leq x \leq 2] \times [-1 \leq y \leq 3]$. (Tip: Use the `plot_surface` function.) Can you guess from the plot if numerical minimization with gradient descent will be fast or slow for this function?

<!-- grade_id: cell-5dfbbae25832f037, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-d2b407ddc3fbb9fc, points: 0, grade: false, solution: true -->
```python
"""
Create function h.
"""
def h(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()

"""
Declare x and y.
"""    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e9cba008bca6b37a, points: 2, grade: true, solution: false -->
```python
"""
Hidden test for function h.
"""
```

<!-- grade_id: cell-85ead8c8bd53e79c, points: 2, grade: true, solution: true -->
```python
"""
Create a function to plot h.
"""
def plot_h(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()

plot_h(x,y)
```

### Exercise 2.2
Knowing that a critical point of a function is a point where the gradient vanishes, show that $(1, 1)$ is the unique critical point of $h$.  Prove that this point is a minimum for $h$. 

<!-- grade_id: cell-1336cb1e9d7d38bd, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 2.3
Write down the gradient descent iteration rule for function $h$. 

<!-- grade_id: cell-91163736650b6698, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 2.4
Implement gradient descent. Try some different values of $\eta$. Does the algorithm converge? How fast? Make plots of the trajectories on top of a contour plot of $h$. (Hint: have a look at the example contour_example.py on Brightspace for inspiration to plot contours of functions and trajectories). Report your findings. Explain why numerical minimization with gradient descent is slow for this function.

First implement the derivative of $h(x,y)$.

<!-- grade_id: cell-ebfab8f71f837375, points: 0, grade: false, solution: true -->
```python
def dh_dxy(x, y):
    """
    This function is the derivative of the function h(x, y).
    
    Parameters
    ----------
    x : float
        data point from x-axis
    y : float
        data point from y-axis
    
    Returns
    -------
    vals : array
        NumPy array of parameter values computed during minimization
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-14c1d4dd062d9e68, points: 1, grade: true, solution: false -->
```python
"""
Test for dh_dxy.
"""
assert np.array_equal(dh_dxy(1, 1), np.array([0, 0]))
```

Now implement the gradient descent algorithm.

<!-- grade_id: cell-e5e7a55f7e7d740f, points: 0, grade: false, solution: true -->
```python
def grad_descent(grad, val_init, eta, max_iter, tol):
    """ This function implements the gradient descent algorithm.
    
    Parameters
    ----------
    grad : function
        Returns the derivative of the function with respect to the pair (x, y).
    val_init : tuple
        Initial values for parameters
    eta : float
        Gradient descent learning rate
    max_iter : int
        Maximum number of gradient descent iterations
    tol : float
        Tolerance for detecting convergence
    
    Returns
    -------
    vals : array
        NumPy array of parameter values computed during minimization
    dists : array
        NumPy array of distances from the current point to the previous point
    tot_iter : int
        Number of performed gradient descent iterations
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9cce6b818f4aa52e, points: 3, grade: true, solution: false -->
```python
"""
Hidden test for grad_descent.
"""
```

Finally, run the gradient descent algorithm with different values of $\eta$.

<!-- grade_id: cell-76876feddeb4a5f0, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Explain what you see!

<!-- grade_id: cell-40c451de6dac1163, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE