1. **Restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart)
2. **Run all cells** (in the menubar, select Cell$\rightarrow$Run All).
3. __Use the__ `Validate` __button in the Assignments tab before submitting__.

__Include comments, derivations, explanations, graphs, etc.__ 

You __work in groups__ (= 3 people). __Write the full name and S/U-number of all team members!__

GROUP NUMBER (brightspace): 13
* Andrew Schroeder, s1111686
* Stian Grønlund,   s1122151
* Monika Kazlauskaite, s1126411

---

# Assignment 3 (Statistical Machine Learning 2024)
# **Deadline: 22 November 2024**

## Instructions
* Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE` __including comments, derivations, explanations, graphs, etc.__ 
Elements and/or intermediate steps required to derive the answer have to be in the report. If an exercise requires coding, explain briefly what the code does (in comments). All figures should have titles (descriptions), axis labels, and legends.
* __Please use LaTeX to write down equations/derivations/other math__! How to do that in Markdown cells can be found [here](https://www.fabriziomusacchio.com/blog/2021-08-10-How_to_use_LaTeX_in_Markdown/), a starting point for various symbols is [here](https://www.overleaf.com/learn/latex/Mathematical_expressions).
* Please do __not add new cells__ to the notebook, try to write the answers only in the provided cells. Before you turn the assignment in, make sure everything runs as expected.
* __Use the variable names given in the exercises__, do not assign your own variable names. 
* __Only one team member needs to upload the solutions__. This can be done under the Assignments tab, where you fetched the assignments, and where you can also validate your submissions. Please do not change the filenames of the individual Jupyter notebooks.

For any problems or questions regarding the assignments, ask during the tutorial or send an email to charlotte.cambiervannooten@ru.nl and janneke.verbeek@ru.nl .

## Introduction
Assignment 3 consists of:
1. The faulty lighthouse (30 points);
2. __Gaussian processes (40 points)__;
3. Bayesian polynomial regression (30 points).

## Libraries

Please __avoid installing new packages__, unless really necessary.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it to at least version 3."

import numpy as np
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import Axes3D

# Set fixed random seed for reproducibility
np.random.seed(2022)
```

## Gaussian processes (40 points)
### Part 1 - Sampling from Gaussian stochastic processes
One widely used kernel function for Gaussian process regression is given by the exponential of a quadratic form, with the addition of constant and linear terms (eq. 6.63 Bishop):
\begin{equation}
    k(\mathbf{x}, \mathbf{x'}) = \theta_0 \exp\left( - \frac{\theta_1}{2} || \mathbf{x} - \mathbf{x'} ||^2 \right) + \theta_2 + \theta_3 \mathbf{x}^T \mathbf{x'}
    \label{eqn:kernel}
    \tag{1}
\end{equation} 
We denote by $\boldsymbol{\theta} = (\theta_0, \theta_1, \theta_2, \theta_3)$ the hyperparameter vector governing the kernel function $k$. 
1. Implement the kernel given by equation \eqref{eqn:kernel} as a function of $\mathbf{x}$, $\mathbf{x'}$ and $\mathbf{\theta}$. Note that $\mathbf{x}$ can have any dimension.

<!-- grade_id: cell-244cdadcf850dd3a, points: 0, grade: false, solution: true -->
```python
def kernel(x1, x2, theta):
    """
    Define the multivariate kernel function.

    Parameters
    ----------
    x1 : (D,) np.ndarray
        A D-dimensional data point.
    x2 : (D,) np.ndarray
        Another D-dimensional data point.
    theta : (4,) np.ndarray
        The array containing the hyperparameters governing the kernel function.

    Returns
    -------
    float
        Value of the kernel function.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-350de9922d852045, points: 3, grade: true, solution: false -->
```python
"""
Hidden test for the kernel function.
"""
```

2. We first consider the univariate case, for the parameter values $\mathbf{\theta} = (1, 1, 1, 1)$ and $N = 101$ equally spaced points $\mathbf{X}$ in the interval $[-1, 1]$. Compute the Gram matrix $\mathbf{K}(\mathbf{X}, \mathbf{X})$ (eq. 6.54 Bishop).

__Note:__ Even when $\mathbf{K}$ is positive (semi)definite, some of its eigenvalues may be too small to accurately compute (same for the determinant). This may pose a problem when generating a multivariate Gaussian distribution using $\mathbf{K}$ as its covariance matrix. You can alleviate this issue by adding a very small (negligible) diagonal term to $\mathbf{K}$. 

<!-- grade_id: cell-15b47d6adb1aa8cd, points: 0, grade: false, solution: true -->
```python
"""
Define the parameter values.

Parameters
----------
theta : (4,) np.ndarray
    Hyperparameters for the kernel function.
X : (N,) np.ndarray
    Array of equally spaced points.
"""
theta = np.array([1, 1, 1, 1])
N = 101
X = np.linspace(-1, 1, N)

def compute_Gram_matrix(X, theta):
    """
    Compute the Gram matrix.

    Parameters
    ----------
    X : (N,D) np.ndarray
        The array containing the data points.
    theta : (4,) np.ndarray
        The array containing the hyperparameters governing the kernel function.

    Returns
    -------
    Gram matrix. (np.ndarray)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-de2dd2da1f174340, points: 4, grade: true, solution: false -->
```python
"""
Tests for compute_Gram_matrix.
"""
K = compute_Gram_matrix(X, theta)
assert type(K) is np.ndarray
```

What is the size of $\mathbf{K}$?

<!-- grade_id: cell-41b2b53ef1db9a09, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

3. We will now use the previously computed matrix $\mathbf{K}(\mathbf{X}, \mathbf{X})$ to produce samples from the Gaussian process prior $\mathbf{y}(\mathbf{X}) \sim \mathcal{N}(\mathbf{0}, \mathbf{K}(\mathbf{X}, \mathbf{X}))$, with $\mathbf{X}$ being the previously determined $N$ equally spaced points. Generate five functions $\mathbf{y}(\mathbf{X})$ and plot them against the $N$ input values $\mathbf{X}$. Repeat this process (remember to compute a new $\mathbf{K}$ each time) for the hyperparameter configurations from Bishop, Figure 6.5:
\begin{equation}
    \mathbf{\theta} \in \{ (1, 4, 0, 0), (9, 4, 0, 0), (1, 64, 0, 0), (1, 0.25, 0, 0), (1, 4, 10, 0), (1, 4, 0, 5) \}.
\end{equation} 

<!-- grade_id: cell-91a8983f8779941c, points: 5, grade: true, solution: true -->
```python
def generate_curves_GP_prior(theta, num_samples = 5):
    """
    Generate univariate function samples (curves) from the Gaussian process prior.

    Parameters
    ----------
    theta : (4,) np.ndarray
        The array containing the hyperparameters governing the kernel function.
    num_samples : integer
        The number of functions to plot.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7e568c52a1390551, points: 2, grade: true, solution: true -->
```python
"""
Plot curves with different values of theta.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

Describe the differences between the plots. Explain in which way each of the kernel parameters affects the generated samples.

<!-- grade_id: cell-9bf4c95c7a2f9c37, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

4. We now move to the bivariate case. Instead of an interval, we now consider a 2-D grid of equally spaced points of size $N = 21 \times 21$ in $[-1, 1] \times [-1, 1]$. We collect all these grid points in a data matrix $\mathbf{X}$, where each one of the 441 observations has two dimensions.

<!-- grade_id: cell-b91ef28dcf1e7306, points: 0, grade: false, solution: true -->
```python
"""
Compute the Gram matrix in the bivariate case.

Parameters
----------
K : np.ndarray
    The Gram matrix.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2f4ded17a4d217db, points: 3, grade: true, solution: false -->
```python
"""
Tests for bivariate Gram matrix.
"""
```

What is the size of $\mathbf{K}$ now? What does this tell you about the scalability of sampling multivariate functions from Gaussian processes in higher dimensions?

<!-- grade_id: cell-478c45187cb376ee, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

5. Using the same kernel from equation \eqref{eqn:kernel}, compute the Gram matrix $\mathbf{K}(\mathbf{X}, \mathbf{X})$ on the grid for each hyperparameter configuration $\mathbf{\theta} \in \{ (1, 1, 1, 1), (1, 10, 1, 1), (1, 1, 1, 10) \}$. For each $\mathbf{K}$, generate and plot five random surfaces from each Gaussian process prior  $\mathbf{y}(\mathbf{X}) \sim \mathcal{N}(\mathbf{0}, \mathbf{K}(\mathbf{X}, \mathbf{X}))$.

<!-- grade_id: cell-e6b906b954c3a1c4, points: 4, grade: true, solution: true -->
```python
def generate_surfaces_GP_prior(theta, num_samples = 5):
    """
    Generate bivariate function samples (surfaces) from the Gaussian process prior.

    Parameters
    ----------
    theta : (4, ) array
        The array containing the hyperparameters governing the kernel function.
    num_samples : integer
        The number of functions to plot.

    Returns
    -------
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fe6416cb0e5fdc58, points: 2, grade: true, solution: true -->
```python
"""
Plot curves with different values of theta.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

Compare the observed differences to the univariate case.

<!-- grade_id: cell-2c7643916c9f098c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Part 2 - Gaussian processes for regression
We would like to apply Gaussian process models to the problem of regression (Bishop 6.4.2). We consider a noisy model of the form:
\begin{equation*}
t_n = y_n + \epsilon_n,
\end{equation*}
where $y_n = y(\mathbf{x}_n)$ and $\epsilon_n$ are i.i.d. samples from a random noise variable on the observed target values. Furthermore, we assume that the noise process has a Gaussian distribution given by:
\begin{equation}
    p(t_n | y_n) = \mathcal{N}(t_n | y_n, \beta^{-1})
    \label{eqn:gp_noise_process}
    \tag{2}
\end{equation}

Going back to a one-dimensional input space, we consider the following training data consisting of four data points:
\begin{equation*}
	\mathcal{D} = \{ (x_1 = -0.5, t_1 = 0.5), (x_2 = 0.2, t_2 = -1), (x_3 = 0.3, t_3 = 3), (x_4 = -0.1, t_4 = -2.5)\}.
\end{equation*}
1. Just as before, compute the Gram matrix of the training data for $\mathbf{\theta} = (1, 1, 1, 1)$. Then, taking $\beta = 1$ in equation \eqref{eqn:gp_noise_process}, compute the covariance matrix $\mathbf{C}$ corresponding to the marginal distribution of the training target values: $p(\mathbf{t}) = \mathcal{N}(\mathbf{t} | \mathbf{0}, \mathbf{C})$.

<!-- grade_id: cell-02cb6c94ecd75584, points: 0, grade: false, solution: true -->
```python
"""
Compute the marginal distribution covariance.

Parameters
----------
C : (4,4) np.ndarray
    The covariance matrix corresponding to the marginal distribution.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8698de2d3e9b656a, points: 3, grade: true, solution: false -->
```python
"""
Tests to check the covariance matrix C.
"""
assert type(C) is np.ndarray
assert C.shape == (4,4)
```

2. Using the previous results, compute the mean and the covariance of the conditional distribution $p(t | \mathbf{t})$ of a new target value $t$ corresponding to the input $x = 0$. Which equations from Bishop do you need? Write them down here.

<!-- grade_id: cell-501fa135983664b3, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Compute the mean and the covariance in the cell below. 

<!-- grade_id: cell-829193a0edec6858, points: 0, grade: false, solution: true -->
```python
"""
Compute the conditional predictive distribution.

Parameters
----------
m : float
    The mean of the conditional predictive distribution.
s2 : float
    The (co)variance of the conditional predictive distribution.
"""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-03872a532232b313, points: 2, grade: true, solution: false -->
```python
"""
Tests to check the mean m and (co)variance s2.
"""
assert (type(m) is float) or (type(m) is np.float64)
assert (type(s2) is float) or (type(s2) is np.float64)
```

3. Does the mean of the conditional distribution $p(t | \mathbf{t})$ go to zero in the limit $x \rightarrow \pm \infty$? If so, explain why this happens. If not, how would you set the parameters $\mathbf{\theta}$ of the kernel function to make it happen?

<!-- grade_id: cell-b5dd1f4b9613f281, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE